<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<!-- question: 111843  -->
  <question type="stack">
    <name>
      <text>RREF - interactive (2)</text>
    </name>
    <questiontext format="plain_text">
      <text><![CDATA[<script>

hide_all_parts = function(question_name) {

console.log(question_name);

    var parts_to_hide = document.querySelectorAll('div[id*="'+question_name+'_part"]');

    parts_to_hide.forEach((qpart) => {

      qpart.style.display = "none";

    });

}

show_part = function(part_to_show) {

    console.log("Show part: "+part_to_show);

    document.addEventListener("DOMContentLoaded", function(){

    var node_to_show = document.querySelector('div[id*="'+part_to_show+'"]');

    node_to_show.style.display = "block";

    document.querySelector('div[id=RREF_visible_parts]').appendChild(node_to_show);

  });

};

document.addEventListener("DOMContentLoaded", function(){

  hide_all_parts("RREF");

});

show_part("RREF_part1");



</script>


<div id="RREF">



<div id="RREF_visible_parts"></div>

<div id="RREF_part1">

<p>Give an example of 3x3 matrix in RREF of rank 3 </p>

<p>[[input:ans1]] [[validation:ans1]] [[feedback:prt1]]</p>

</div>



<div id="RREF_part2">

<p>Give an example of 3x3 matrix in RREF of rank 2 </p>

<p>[[input:ans2]] [[validation:ans2]] [[feedback:prt2]]</p>

</div>



<div id="RREF_part3">

<p>Give an example of 3x3 matrix in RREF of rank 1 </p>

<p>[[input:ans3]] [[validation:ans3]] [[feedback:prt3]]</p>

</div>



<div id="RREF_part4">

<p>Give an example of 3x3 matrix in RREF of rank 0 </p>

<p>[[input:ans4]] [[validation:ans4]] [[feedback:prt4]]</p>

</div>


<div id="RREF_part5">

<p>Now give a matrix in RREF where at least one entry is not 1 or 0 </p>

<p>[[input:ans5]] [[validation:ans5]] [[feedback:prt5]]</p>

</div>


</div>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <stackversion>
      <text>2022082900</text>
    </stackversion>
    <questionvariables>
      <text>rref1: matrix([1,0,0],[0,1,0],[0,0,1]);
rref2: matrix([0,1,0],[0,0,1],[0,0,0]);
rref3: matrix([0,0,1],[0,0,0],[0,0,0]);
rref4: matrix([0,0,0],[0,0,0],[0,0,0]);
rref5: matrix([1,0,3/5],[0,1,-3],[0,0,0]);
is_rref_not01(answer):= is(ans5[2]=[0,0,0] and ans5[3] = [0,0,0] and ((ans5[1,1] = 1 and notequal(ans5[1,2],0) and notequal(ans5[1,2],1) and notequal(ans5[1,3],0) and notequal(ans5[1,3],1)) or (ans5[1,1] = 0 and ans5[1,2] = 1)) or ans5[3] = [0,0,0] and ((ans5[2] = [0,0,1] and ans5[1,1] = 1 and ans5[1,3] = 0 and notequal(ans5[1,2],0) and notequal(ans5[1,2],1)) or (notequal(ans5[2,3],0) and notequal(ans5[2,3],1) and notequal(ans5[1,3],0) and notequal(ans5[1,3],1) and ans5[1,1] = 1 and ans5[1,2] = 0 and ans5[2,1] = 0 and ans5[2,2] = 1)));</text>
    </questionvariables>
    <specificfeedback format="html">
      <text></text>
    </specificfeedback>
    <questionnote>
      <text></text>
    </questionnote>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <variantsselectionseed></variantsselectionseed>
    <input>
      <name>ans1</name>
      <type>matrix</type>
      <tans>rref1</tans>
      <boxsize>3</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint></syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords></forbidwords>
      <allowwords></allowwords>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options></options>
    </input>
    <input>
      <name>ans2</name>
      <type>matrix</type>
      <tans>matrix([1,0,0],[0,1,0],[0,0,0])</tans>
      <boxsize>3</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint></syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords></forbidwords>
      <allowwords></allowwords>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options></options>
    </input>
    <input>
      <name>ans3</name>
      <type>matrix</type>
      <tans>rref3</tans>
      <boxsize>3</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint></syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords></forbidwords>
      <allowwords></allowwords>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options></options>
    </input>
    <input>
      <name>ans4</name>
      <type>matrix</type>
      <tans>rref4</tans>
      <boxsize>3</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint></syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords></forbidwords>
      <allowwords></allowwords>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options></options>
    </input>
    <input>
      <name>ans5</name>
      <type>matrix</type>
      <tans>rref5</tans>
      <boxsize>3</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint></syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords></forbidwords>
      <allowwords></allowwords>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>1</showvalidation>
      <options></options>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text>is_rref: is(ans1 = rref1)</text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <answertest>AlgEquiv</answertest>
        <sans>is_rref</sans>
        <tans>true</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="plain_text">
          <text><![CDATA[<script>
    show_part("RREF_part2");
</script>]]></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p dir="ltr" style="text-align: left;">That's not a 3x3 RREF matrix of rank 3.&nbsp;Try again!&nbsp;<br></p>]]></text>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt2</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text>is_rref: false
if (transpose(ans2)[1] = [1,0,0] and transpose(ans2)[2] = [0,1,0] and ans2[3,3] = 0) or (transpose(ans2)[1] = [1,0,0] and transpose(ans2)[3] = [0,1,0] and ans2[2,2] = 0 and ans2[2,3] = 0) or (ans2 = rref2)
    then is_rref: true;</text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <answertest>AlgEquiv</answertest>
        <sans>is_rref</sans>
        <tans>true</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-1-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<script>
    show_part("RREF_part3");
</script>]]></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt2-1-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p dir="ltr" style="text-align: left;">That's not a 3x3 RREF matrix of rank 2.&nbsp;Try again!&nbsp;<br></p>]]></text>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt3</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text>is_rref: false
if (ans3[2] = [0,0,0] and ans3[3] = [0,0,0] and (ans3[1] = [0,0,1] or ans3[1,1] = 1 or (ans3[1,1] = 0 and ans3[1,2] = 1))) or ans3 = rref3
    then is_rref: true;</text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <answertest>AlgEquiv</answertest>
        <sans>is_rref</sans>
        <tans>true</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt3-1-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<script>
    show_part("RREF_part4");
</script>]]></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt3-1-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p dir="ltr" style="text-align: left;">That's not a 3x3 RREF matrix of rank 1.&nbsp;Try again!&nbsp;</p>]]></text>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt4</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text>is_rref: false
if ans4 = rref4
    then is_rref: true;</text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <answertest>AlgEquiv</answertest>
        <sans>is_rref</sans>
        <tans>true</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty></truepenalty>
        <truenextnode>1</truenextnode>
        <trueanswernote>prt4-1-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<script>
    show_part("RREF_part5");
</script>]]></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt4-1-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p dir="ltr" style="text-align: left;">Not a 3x3 RREF matrix of rank 0!</p>]]></text>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <answertest>AlgEquiv</answertest>
        <sans>all_listp(is_rref_not01,[ans1,ans2,ans3,ans4])</sans>
        <tans>true</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt4-2-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p dir="ltr" style="text-align: left;">All of your inputs have only had 1s and 0s so far.</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt4-2-F</falseanswernote>
        <falsefeedback format="html">
          <text></text>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt5</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text>is_rref_not01(answer):= is(ans5[2]=[0,0,0] and ans5[3] = [0,0,0] and ((ans5[1,1] = 1 and (notequal(ans5[1,2],0) and notequal(ans5[1,2],1)) or (notequal(ans5[1,3],0) and notequal(ans5[1,3],1))) or (ans5[1,1] = 0 and ans5[1,2] = 1)) or ans5[3] = [0,0,0] and ((ans5[2] = [0,0,1] and ans5[1,1] = 1 and ans5[1,3] = 0 and notequal(ans5[1,2],0) and notequal(ans5[1,2],1)) or ((notequal(ans5[2,3],0) and notequal(ans5[2,3],1)) or (notequal(ans5[1,3],0) and notequal(ans5[1,3],1)) and ans5[1,1] = 1 and ans5[1,2] = 0 and ans5[2,1] = 0 and ans5[2,2] = 1)))</text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <answertest>AlgEquiv</answertest>
        <sans>is_rref_not01(ans5)</sans>
        <tans>true</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt5-1-T</trueanswernote>
        <truefeedback format="html">
          <text></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt5-1-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p dir="ltr" style="text-align: left;">Your answer either isn't a 3x3 matrix in RREF or is only made up of 1s and 0s. Try again!</p>]]></text>
        </falsefeedback>
      </node>
    </prt>
  </question>

</quiz>